<a href="https://colab.research.google.com/github/gt-cse-6040/bootcamp_m0s3/blob/main/m0s3nb5_mpe_9.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Problem 9

This problem consists of a single exercise worth ten (10) points.

**Exercise 0** (10 points): Write a funtion, `digitize(n)`, that, given a postive number `n`, returns the digits of `n` as the integer elements of a list in reversed order.

For example:

```python
    assert digitize(1345328) == [8, 2, 3, 5, 4, 3, 1]
```

```python
def digitize(n):
    assert n > 0, '`n` is not a positive number'
    assert type(n) is int, '`n` must be an `int`'
    ### YOUR CODE HERE

#     GOAL
#         Reverse the order of the digits of a number, and return a list of the digits in the reversed order.

#     INPUT
#             A positive number with any number of digits (no digit separators)

#     STRATEGY
#             Note that the requirments of being a positive integer are already done in the asserts
#             Convert the integer to a string
#             Two ways of doing this:
#                  1. Reverse the string
#                  2. Convert the string to a list
#                  3. Convert the individual elements in the list to from strings to integers

#                  a. Convert the string to a list
#                  b. Reverse the list
#                  c.
#             Return the reversed string

#     SOLUTION

    # create the list to return
    return_list = []
    # convert the integer to a string
    str_n = str(n)

    # ****************
    # reverse the string, then make it a list
    # ******************
    # reverse the string
    reverse_str = str_n[::-1]
#     print(reverse_string)
    # convert to list
    reverse_list = list(reverse_str)
#     print(reverse_list)
    # ****************

#     # ****************
#     # make the string a list, then reverse the list
#     # convert the string to a list
#     reverse_list = list(str_n)
#     # reverse the list
#     reverse_list.reverse()
# # #     print(reverse_list)
#     # ****************

#     loop over the list to convert each element to integer
    for i in range(0, len(reverse_list)):
        return_list.append(int(reverse_list[i]))

    return return_list
```

<!-- grade_id: test_digitize, points: 10, grade: true, solution: false -->
```python
# Test_cell: test_digitize

assert digitize(1345328) == [8, 2, 3, 5, 4, 3, 1], 'Wrong output. Make sure your function returns the digits in reverse order'
print("\n(Passed!)")
```

**Fin!** You've reached the end of this problem. Don't forget to restart the
kernel and run the entire notebook from top-to-bottom to make sure you did
everything correctly. If that is working, try submitting this problem. (Recall
that you *must* submit and pass the autograder to get credit for your work!)